```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment_02.ipynb")
```

<!-- BEGIN QUESTION -->

<div class="alert alert-info">

# Submission instructions
rubric={mechanics}

<p>You receive marks for submitting your lab correctly, please follow these instructions:</p>

<ul>
  <!-- <li>Make at least three commits.</li>-->
  <li>Export to HTML via <code>File -> Save and export notebook as... -> HTML</code> and upload the HTML file to Canvas.</li>
    <ul>
      <li>Before submitting, make sure you restart the kernel and rerun all cells.</li>
    </ul>
  <li>Don't change any variable names that are given to you, don't move cells around, and don't include any code to install packages in the notebook.</li>
  </li>
</ul>
</div>

_Points:_ 1

<!-- END QUESTION -->

## Warm Up


<div class="alert alert-info">

### Question 0a

Before Proceeding to the main task, here are some important terms and definitions that will be useful for you to remember in order to complete this assignment:
    
<table>
<tr>
    <td> <b>Definitions:</b></td>
<td> <b>Match the terms on the left with the definitions from the list below:</b></td>
</tr>
<tr>
<td> 
    
1. directory
2. command line
3. operating systems
4. root directory
5. current working directory
6. parent directory
7. home directory
8. prompt
9. command line argument
10. absolute path   
</td>
<td> 
        
A) `A directory contained within another directory`  
B) `$`  
C) `A description that specifies the location of a file or directory within a file system.`  
D) `folder`  
E. `/`  
F. `..`  
G. `Mac, Linux, Windows`  
H. `~`  
I. `The directory location in which the program operates`  
J. `-F`  
K. `.`  
L. `shell`  
M. `JupyterLab`  
N. `The part of the operating system that manages how files are stored and retrieved`  
    </td>
    </tr>
</table>

</div>


_Points:_ 1

```python
# Your response should be a sequence of 10 letters, in the order of the corresponding answer, e.g. 'ABCKCCCDFI'
q0a = 'DLGEIFHBJC'
```

```python
grader.check("q0a")
```

Output:
```
q0a results: All test cases passed!
```

## Set Up

1. Download the repo `crime_scene` and save it to your home directory (if you are not sure what your home directory is, type `pwd` in the terminal to see where you are and use `~` to navigate to your home directory).
2. After you saved the folder in your home directory, unzip the folder.
3. Open the terminal and using `cd` command navigate to the folder `crime_scene`


<div class="alert alert-info">
    
### Question 0b
    
What will be your location after step 3 using the relative paths?

A. `"../crime_scene"`  
B. `"./crime_scene"`  
C. `"/crime_scene"`  
D. `"~/crime_scene"`  
    
    
</div>



_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q0b = 'AB'
```

```python
grader.check("q0b")
```

Output:
```
q0b results: All test cases passed!
```

Once your are inside the `crime_scene` directory you can start solving the Bash Murder Mystery.

-----

# Bash Murder Mystery

- For questions where multiple answers are allowed, you need to select all correct answers to be awarded full points (there is always at least one correct answer).
- If there are commands or option flags that you don't recognize, we encourage you to use the manual pages to read about what these
do and try out the commands on your system. Some flags can be used even if they aren't necessarily needed, but won't change the code outcome.
- Unless the instructions explicitly tell you to change/navigate to another directory,
you are expected to perform the task from the directory you are currently in
- Instructions for this mystery murder game do not explicitly mention the commands that you need to use, so try to think about what commands can be used to complete each step. If you get stuck, check the hints!

**Step 1: Enter the crime scene.**

<details><summary><b>Hint:</b></summary>
Use the <code>cd</code> command to enter the <code>crime_scene</code> folder and the <code>ls -F</code> command to find out which directories there are inside.
</details>

<br>



<div class="alert alert-info">
    
### Question 1
    
What *folders* do you find at the crime scene?

A. `Suspects`  
B. `victim.txt`  
C. `note_from_killer.txt`  
D. `Backup`  
E. `security_camera_footage`

</div>


_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q1 = 'AE'
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

------

**Step 2: Open the victim record and add your name as an investigator**

<details><summary><b>Hint:</b></summary>
Use the <code>code</code> command to open the victim file and enter your name on line 8 of the text file, then save the file.
</details>




<div class="alert alert-info">
    
### Question 2
    
What was the victim's time of death?

A. `Not available`  
B. `8:30 AM`  
C. `7:30 PM`  
D. `8:30 PM`  
E. `10:30 PM` 

</div>


_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q2 = 'D'
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

------

**Step 3: Find the hidden murder weapon**

<details><summary><b>Hint 1:</b></summary>
What <code>ls</code> command can you use to see hidden files (check the manual)?
</details>

<details><summary><b>Hint 2:</b></summary>
Use the <code>ls -a</code> command inside the crime_scene to find the murder weapon.
</details>



<div class="alert alert-info">
    
### Question 3
    
Enter the name of the file containing the murder weapon information

</div>



_Points:_ 1

```python
# # Your answer should be a string with name of the file including file extention , e.g., "document.pdf"
q3 = '.murder_weapon.txt'
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

------

**Step 4: Investigate the murder weapon**

<details><summary><b>Hint 1:</b></summary>
Explore the contents of <code>.murder_weapon.txt</code> file
</details>

<details><summary><b>Hint 2:</b></summary>
Use the <code>head</code> command to read the contents of <code>.murder_weapon.txt</code> file.
</details>



<div class="alert alert-info">
    
### Question 4
    
Enter the type of the weapon found at the crime scene:

A. `Weapon was not found`  
B. `Gun`  
C. `Kitchen Knife`  
D. `Rifle`  

</div>


_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q4 = 'C'
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

------

**Step 5: Find a list of the suspects' alibis related to this murder (you may want to explore the contents of the crime scene).**

<details><summary><b>Hint:</b></summary>
Use the <code>cd</code> command to go to Suspects and use <code>ls -F</code> command to check for suspects
</details>


<!-- BEGIN QUESTION -->

<div class="alert alert-info">
    
### Question 5a
    
How many suspects do you have?"

</div>

_Points:_ 1


<!-- END QUESTION -->

_Points:_ 1

```python
# Your answer should be a string with a number that you think are correct, e.g. '11'
q5a = '5'
```

```python
grader.check("q5a")
```

Output:
```
q5a results: All test cases passed!
```


<div class="alert alert-info">
    
### Question 5b
    
Navigate back to parent directory (`crime scene` folder). 

You can also check the number of suspects using a more complex query from our notes file, combining
the power of  `ls`  and  `wc -l`  with a pipe. 
    
Type your solution below:

</div>


_Points:_ 1

```python
# Your answer should be a string with a single line command surrounded by "" 
q5b = "ls Suspects | wc -l"
```

```python
grader.check("q5b")
```

Output:
```
q5b results: All test cases passed!
```

------

**Step 6: Review the alibi of each suspect.**

<details><summary><b>Hint:</b></summary>
Use <code>cd</code> to navigate back to the <code>Suspects</code> directory, and use the <code>head</code> command to read the first 10 lines of each alibi
</details>





<div class="alert alert-info">
    
### Question 6a 
    
Which command(s) can you use to read each suspects' **entire** alibi file? Assume the file is called `filename.md`

A. `head filename.md`   
B. `code filename.md`  
C. `code .`   
D. `cat filename.md`  
E. `head -n 25 filename.md`  
F. `rm`    
G. `mkdir`    


</div>



_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q6a = 'BDE'
```

```python
grader.check("q6a")
```

Output:
```
q6a results: All test cases passed!
```


<div class="alert alert-info">
    
### Question 6b
    
You most likely you checked the alibi of each suspect individually. However, this repetitive task can be simplified using the wildcards. Which command can you use to do this more quickly?

A. `head *.txt`  
B. `cd --all`  
C. `head .txt`  
D. `head -n 25 *.txt`  
E. `head --all`  
F. `head ?`  


</div>



_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q6b = 'D'
```

```python
grader.check("q6b")
```

Output:
```
q6b results: All test cases passed!
```

------

**Step 7: Find a suspect who was at the "crime scene" during the murder.**

<details><summary><b>Hint:</b></summary>
Use the <code>grep</code> command together with a wildcard to search all the files.
</details>



<div class="alert alert-info">
    
### Question 7
    
Enter the command that will allow you to check which suspect been at the "crime scene". Use one command.

</div>


_Points:_ 1

```python
q7 = 'grep "crime scene" *.txt'
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

------

**Step 8: Return to the crime scene and explore the content from the security cameras.**

<details><summary><b>Hint:</b></summary>
Use the <code>cd</code> command to get back to the crime scene and to move into the security cameras directory
</details>


<!-- BEGIN QUESTION -->

<div class="alert alert-info">
    
### Question 8a
    
What command can you use to get to security camera footage directly from Suspects folder?

A. `cd ./security_camera_footage`  
B. `cd ../security_camera_footage`  
C. `cd ~/security_camera_footage`  
D. `cd crime_scene/security_camera_footage`  
    

</div>

_Points:_ 1


<!-- END QUESTION -->

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q8a = 'B'
```

```python
grader.check("q8a")
```

Output:
```
q8a results: All test cases passed!
```


<div class="alert alert-info">
    
### Question 8b
    
Who was at the crime scene on the day of the murder?

A. Aada Aalto  
B. Kofi Mensah  
C. Ye Wong  
D. Deepika Sharma  
E. Emily Davis

</div>



_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q8b = 'AE'
```

```python
grader.check("q8b")
```

Output:
```
q8b results: All test cases passed!
```

------

**Step 9: Based on each suspect's alibi and video camera footage, remove the alibi of the suspect who was at the crime scene at the time of murder**

<details><summary><b>Hint:</b></summary>
Use the <code>rm</code> command to remove the alibi of the suspect who was at the crime scene at the time of murder
</details>
<br>


<div class="alert alert-info">
    
### Question 9a
    
What command can you use to return to a crime scene?

A. `cd .`  
B. `cd ~`  
C. `cd ..`  
D. `cd crime_scene`  

</div>


_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q9a = 'C'
```

```python
grader.check("q9a")
```

Output:
```
q9a results: All test cases passed!
```


<div class="alert alert-info">
    
### Question 9b
    
Who do you suspect could have committed the murder? Note - make sure you get this question right before proceeding to next step

A. Aada Aalto  
B. Kofi Mensah  
C. Ye Wong  
D. Deepika Sharma  
E. Emily Davis

</div>


_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q9b = 'E'
```

```python
grader.check("q9b")
```

Output:
```
q9b results: All test cases passed!
```


<div class="alert alert-info">
    
### Question 9c
    
Which command can you use to remove the alibi file of the suspect who lied about when they were at the crime scene (given that you are currently in the `security_camera_footage` folder)? 

We don't want to give away who the murderer is yet, so we have replaced the suspects' alibi numbers with `X`, you should replace it with the number of the suspect that you think is the murderer.

A. `rm *alibi_suspectX.txt`  
B. `rm ../Suspects/alibi_suspectX.txt`  
C. `rm /Suspects/alibi_suspectX.txt`  
D. `rm Suspects/alibi_suspectX.txt`  

</div>


_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q9c = 'B'
```

```python
grader.check("q9c")
```

Output:
```
q9c results: All test cases passed!
```


------

**Step 10: Check if the alibi file is removed.**

<details><summary><b>Hint:</b></summary>
Use the <code>ls</code> command to check if the alibi file is removed.
</details>







<div class="alert alert-info">
    
### Question 10
    
How can you check that the alibi file has been removed? Type out the command to do this from your current directory (still `security_camera_footage`).


</div>


_Points:_ 1

```python
# Your answer should be a string with a single line command surrounded by "" 
q10 = 'ls ../Suspects'
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

------

**Step 11: Create a report (aka directory) called Backup.**

<details><summary><b>Hint:</b></summary>
Use the <code>mkdir</code> command to create a <code>backup</code> directory
</details>
 


<div class="alert alert-info">
    
### Question 11
    
Which of the following commands could be used to to create a new directory?

A. `code backup`  
B. `cd backup`  
C. `makedir -p backup`  
D. `mkdir -p backup`  
E. `makedir backup`  
F. `mkdir backup`   
G. `cd --create backup`    

</div>


_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q11 = 'F'
```

```python
grader.check("q11")
```

Output:
```
q11 results: All test cases passed!
```

------

**Step 12: Move the remaining alibis to backup directory.** 

<details><summary><b>Hint:</b></summary>
Use the <code>mv</code> command to move the remaining alibis to Backup directory
</details>



<div class="alert alert-info">
    
### Question 12
    
State you final decision. Who was the murderer?

A. Aada Aalto  
B. Kofi Mensah  
C. Ye Wong  
D. Deepika Sharma  
E. Emily Davis

</div>


_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q12 = 'E'
```

```python
grader.check("q12")
```

Output:
```
q12 results: All test cases passed!
```

<div class="alert alert-danger" style="color:black">
    
**Restart and run all before submitting**
    
Before submitting,
don't forget to run all cells in your notebook
to make sure there are no errors
and so that the TAs can see the output of all the cells properly.
You can do this by clicking the ▶▶ button
or going to `Kernel -> Restart Kernel and Run All Cells...` in the menu.
This is not only important for this course,
but a good habit you should get into before ever committing a notebook to GitHub,
so that your collaborators can run it from top to bottom
without issues.
</div>

------

Mystery Solved, Congratulations! 👍




---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q0a results: All test cases passed!

q0b results: All test cases passed!

q1 results: All test cases passed!

q10 results: All test cases passed!

q11 results: All test cases passed!

q12 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!

q5a results: All test cases passed!

q5b results: All test cases passed!

q6a results: All test cases passed!

q6b results: All test cases passed!

q7 results: All test cases passed!

q8a results: All test cases passed!

q8b results: All test cases passed!

q9a results: All test cases passed!

q9b results: All test cases passed!

q9c results: All test cases passed!
```

Tests for q0a (points: 1)
```
>>> assert 'b' in q0a.lower()
>>> assert 'c' in q0a.lower()
>>> assert 'd' in q0a.lower()
>>> assert 'e' in q0a.lower()
>>> assert 'f' in q0a.lower()
>>> assert 'g' in q0a.lower()
>>> assert 'h' in q0a.lower()
>>> assert 'i' in q0a.lower()
>>> assert 'j' in q0a.lower()
>>> assert 'l' in q0a.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'a' not in q0a.lower()
>>> assert 'k' not in q0a.lower()
>>> assert 'm' not in q0a.lower()
>>> assert 'n' not in q0a.lower()
```
failure_message: At least one of the options you have selected is not correct
```
>>> 
>>> assert q0a.upper() == 'DLGEIFHBJC'
```
failure_message: At least one of the options you have selected is not in correct order


Tests for q0b (points: 1)
```
>>> assert 'a' in q0b.lower()
>>> assert 'b' in q0b.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers
```
>>> assert 'c' not in q0b.lower()
>>> assert 'd' not in q0b.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q1 (points: 1)
```
>>> assert 'a' in q1.lower()
>>> assert 'e' in q1.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'c' not in q1.lower()
>>> assert 'd' not in q1.lower()
>>> assert 'b' not in q1.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q10 (points: 1)
```
>>> q10_string = q10.replace(" ", "")
>>> assert q10_string == "ls../Suspects"
```
failure_message: Your answer is incorrect


Tests for q11 (points: 1)
```
>>> assert 'f' == q11.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q12 (points: 1)
```
>>> assert 'e' in q12.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q2 (points: 1)
```
>>> assert 'd' == q2.lower()
>>>
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q3 (points: 1)
```
>>> q3_string = q3.replace(" ", "")
>>> assert q3_string == ".murder_weapon.txt"
```
failure_message: At least one of the options you have selected is not in correct order
```
>>> q3_string = q3.replace(" ", "")
>>> assert q3_string != ".murder_weapon"
```
failure_message: file name missing extension
```
>>> q3_string = q3.replace(" ", "")
>>> assert q3_string != "murder_weapon.txt"
```
failure_message: file name is a hidden file. What character at the beginning of this file name?


Tests for q4 (points: 1)
```
>>> assert 'c' == q4.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q5a (points: 1)
```
>>> assert q5a == '5'
```
failure_message: Your answer is incorrect
```
>>> assert q5a != 5
```
failure_message: Your answer should be a string


Tests for q5b (points: 1)
```
>>> q5b_string = q5b.replace(" ", "")
>>> assert q5b_string == "lsSuspects|wc-l"
```
failure_message: Your answer is incorrect


Tests for q6a (points: 1)
```
>>> assert 'b' in q6a.lower()
>>> assert 'd' in q6a.lower()
>>> assert 'e' in q6a.lower()
>>>
```
failure_message: You are missing one of the correct answers
```
>>> assert 'a' not in q6a.lower()
>>> assert 'c' not in q6a.lower()
>>> assert 'f' not in q6a.lower()
>>> assert 'g' not in q6a.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q6b (points: 1)
```
>>> assert 'd' == q6b.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q7 (points: 1)
```
>>> q7_string = q7.replace('"crime scene"', "'crime scene'")
>>> q7_string = q7_string.replace(" ", "")
>>> assert q7_string == "grep'crimescene'*.txt"
```
failure_message: Your answer is incorrect


Tests for q8a (points: 1)
```
>>> assert 'b' == q8a.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q8b (points: 1)
```
>>> assert 'a' in q8b.lower()
>>> assert 'e' in q8b.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'b' not in q8b.lower()
>>> assert 'c' not in q8b.lower()
>>> assert 'd' not in q8b.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q9a (points: 1)
```
>>> assert 'c' == q9a.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q9b (points: 1)
```
>>> assert 'e' == q9b.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q9c (points: 1)
```
>>> assert 'b' == q9c.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers
